# Sprint Challenge
## *Data Science Unit 4 Sprint 1*

After a week of Natural Language Processing, you've learned some cool new stuff: how to process text, how turn text into vectors, and how to model topics from documents. Apply your newly acquired skills to one of the most famous NLP datasets out there: [Yelp](https://www.yelp.com/dataset). As part of the job selection process, some of my friends have been asked to create analysis of this dataset, so I want to empower you to have a head start.  

The real dataset is massive (almost 8 gigs uncompressed). I've sampled the data for you to something more manageable for the Sprint Challenge. You can analyze the full dataset as a stretch goal or after the sprint challenge. As you work on the challenge, I suggest adding notes about your findings and things you want to analyze in the future.

## Challenge Objectives
Successfully complete all these objectives to earn full credit. 

**Successful completion is defined as passing all the unit tests in each objective.**  

Each unit test that you pass is 1 point. 

There are 5 total possible points in this sprint challenge. 


There are more details on each objective further down in the notebook.*
* <a href="#p1">Part 1</a>: Write a function to tokenize the yelp reviews
* <a href="#p2">Part 2</a>: Create a vector representation of those tokens
* <a href="#p3">Part 3</a>: Use your tokens in a classification model on yelp rating
* <a href="#p4">Part 4</a>: Estimate & Interpret a topic model of the Yelp reviews

____

# Before you submit your notebook you must first

1) Restart your notebook's Kernel

2) Run all cells sequentially, from top to bottom, so that cell numbers are sequential numbers (i.e. 1,2,3,4,5...)
- Easiest way to do this is to click on the **Cell** tab at the top of your notebook and select **Run All** from the drop down menu. 

3) Comment out the cell that generates a pyLDAvis visual in objective 4 (see instructions in that section). 
____



### Import Data

<!-- grade_id: cell-395851cd95d17235, points: 0, grade: false, solution: true -->
```python
import pandas as pd

# Load reviews from URL
data_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_4/unit1_nlp/review_sample.json'

# Import data into a DataFrame named df
# YOUR CODE HERE
df = pd.read_json(data_url, lines = True)
```

```python
df
```

Output:
```
business_id  cool                date  funny  \
0     nDuEqIyRc8YKS1q1fX0CZg     1 2015-03-31 16:50:30      0   
1     eMYeEapscbKNqUDCx705hg     0 2015-12-16 05:31:03      0   
2     6Q7-wkCPc1KF75jZLOTcMw     1 2010-06-20 19:14:48      1   
3     k3zrItO4l9hwfLRwHBDc9w     3 2010-07-13 00:33:45      4   
4     6hpfRwGlOzbNv7k5eP9rsQ     1 2018-06-30 02:30:01      0   
...                      ...   ...                 ...    ...   
9995  1h3ysSuSazvXc1aeLiiOew     0 2017-10-07 10:57:15      1   
9996  Rwahe1zbFpw6VIjb5ngZeg     0 2014-01-18 15:52:52      0   
9997  8itGZAOBMiTbHKOwLuh4_Q     0 2018-08-26 02:53:21      0   
9998  A5Rkh7UymKm0_Rxm9K2PJw     0 2018-04-23 23:36:07      0   
9999  xtYiHTmunjfCN2sUaQxBjA     0 2016-06-02 23:37:51      0   

                   review_id  stars  \
0     eZs2tpEJtXPwawvHnHZIgQ      1   
1     DoQDWJsNbU0KL1O29l_Xug      4   
2     DDOdGU7zh56yQHmUnL1idQ      3   
3     LfTMUWnfGFMOfOIyJcwLVA      1   
4     zJSUdI7bJ8PNJAg4lnl_Gg      4   
..
```

<!-- grade_id: cell-cb5006475e42b8f9, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert isinstance(df, pd.DataFrame), 'df is not a DataFrame. Did you import the data into df?'
assert df.shape[0] == 10000, 'DataFrame df has the wrong number of rows.'
```

## Part 1: Tokenize Function
<a id="#p1"></a>

Complete the function `tokenize`. Your function should
- accept one document at a time
- return a list of tokens

You are free to use any method you have learned this week.

```python
# Optional: Consider using spaCy in your function. The spaCy library can be imported by running this cell.
# A pre-trained model (en_core_web_sm) has been made available to you in the CodeGrade container.
# If you DON'T need use the en_core_web_sm model, you can comment it out below.
import spacy
nlp = spacy.load('en_core_web_sm')
```

<!-- grade_id: cell-3d570d5a1cd6cb64, points: 0, grade: false, solution: true -->
```python
def tokenize(doc):
# YOUR CODE HERE
    tokens = []
    for token in nlp(doc):
        if (token.is_stop == False) & (token.is_punct == False):
            tokens.append(token.lemma_)
    return tokens
```

<!-- grade_id: cell-02da164f6fbe730a, points: 1, grade: true, solution: false -->
```python
'''Testing'''
assert isinstance(tokenize(df.sample(n=1)["text"].iloc[0]), list), "Make sure your tokenizer function accepts a single document and returns a list of tokens!"
```

## Part 2: Vector Representation
<a id="#p2"></a>
1. Create a vector representation of the reviews (i.e. create a doc-term matrix).
2. Write a fake review and query for the 10 most similar reviews, print the text of the reviews. Do you notice any patterns?
    - Given the size of the dataset, use `NearestNeighbors` model for this. 

<!-- grade_id: cell-0e96491cb529202c, points: 0, grade: false, solution: true -->
```python
%%time
# Create a vector representation of the reviews 
# Name that doc-term matrix "dtm"

# YOUR CODE HERE
from sklearn.feature_extraction.text import TfidfVectorizer

tfid_vect = TfidfVectorizer(stop_words="english", max_features=1000, ngram_range=(1,2))

dtm = tfid_vect.fit_transform(df['text'])

dtm = pd.DataFrame(data=dtm.toarray(), columns=tfid_vect.get_feature_names())
```

Output:
```
Wall time: 2.9 s
```

<!-- grade_id: cell-3d5bc610a8ec6b24, points: 0, grade: false, solution: true -->
```python
# Create and fit a NearestNeighbors model named "nn"
from sklearn.neighbors import NearestNeighbors

# YOUR CODE HERE
# Fit on DTM
nn = NearestNeighbors(n_neighbors=10, algorithm='kd_tree')
nn.fit(dtm)

```

Output:
```
NearestNeighbors(algorithm='kd_tree', leaf_size=30, metric='minkowski',
                 metric_params=None, n_jobs=None, n_neighbors=10, p=2,
                 radius=1.0)
```

<!-- grade_id: cell-c43704dcff67e99b, points: 1, grade: true, solution: false -->
```python
'''Testing.'''
assert nn.__module__ == 'sklearn.neighbors._unsupervised', ' nn is not a NearestNeighbors instance.'
assert nn.n_neighbors == 10, 'nn has the wrong value for n_neighbors'
```

<!-- grade_id: cell-496203e8746296ca, points: 0, grade: false, solution: true -->
```python
# Create a fake review and find the 10 most similar reviews

# YOUR CODE HERE
# sample a doc from dtm to use as our query point
doc = tfid_vect.transform(["resting 5 minutes for the time they said they close, i arrived to the store and one of the employees close the door on my face, it was so rude sice he looked me and did not care at all, other days they make me wait too long in line and gave me the wrong items, i would never recommend this place to nobody its unfair the treat people like that"])
doc = doc.toarray()

# Query Using kneighbors 
neigh_dist, neigh_index = nn.kneighbors(doc)
```

```python
print(neigh_dist, neigh_index)
```

Output:
```
[[1.         1.         1.16355241 1.16888305 1.20413044 1.20717571
  1.21331433 1.21331919 1.21387337 1.21572704]] [[6204 6311 3787 2859  955 4185 9511 7715 4752 7990]]
```

## Part 3: Classification
<a id="#p3"></a>
Your goal in this section will be to predict `stars` from the review dataset. 

1. Create a pipeline object with a sklearn `CountVectorizer` or `TfidfVector` and any sklearn classifier.
    - Use that pipeline to train a model to predict the `stars` feature (i.e. the labels). 
    - Use that Pipeline to predict a star rating for your fake review from Part 2. 



2. Create a parameter dict including `one parameter for the vectorizer` and `one parameter for the model`. 
    - Include 2 possible values for each parameter
    - **Use `n_jobs` = 1** 
    - Due to limited computational resources on CodeGrader `DO NOT INCLUDE ADDITIONAL PARAMETERS OR VALUES PLEASE.`
    
    
3. Train the entire pipeline with a GridSearch
    - Name your GridSearch object as `gs`

<!-- grade_id: cell-e2beb0252d274bba, points: 0, grade: false, solution: true -->
```python
from sklearn.pipeline import Pipeline
from sklearn.feature_extraction.text import TfidfVectorizer
from sklearn.neighbors import KNeighborsClassifier
from sklearn.ensemble import RandomForestClassifier
from sklearn.model_selection import GridSearchCV

# Name the gridsearch instance "gs"

tfid_vect = TfidfVectorizer(stop_words="english", max_features=1000, ngram_range=(1,2))
clf = RandomForestClassifier()

pipe = Pipeline([
                ('tfid_vect',tfid_vect),
                ('clf', clf)
                ])
y = df['stars']
X = df['text']

# Parameters to search in dictionary 
parameters = {
    'tfid_vect__max_df': (0.75, 1.0),
    'clf__max_depth':(20, 25)
}

# Implement a grid search with cross-validation
gs = GridSearchCV(pipe, param_grid=parameters, n_jobs=1, cv=3, verbose=1)
gs.fit(X,y)
```

Output:
```
Fitting 3 folds for each of 4 candidates, totalling 12 fits
[Parallel(n_jobs=1)]: Using backend SequentialBackend with 1 concurrent workers.
[Parallel(n_jobs=1)]: Done  12 out of  12 | elapsed:   47.6s finished
GridSearchCV(cv=3, error_score=nan,
             estimator=Pipeline(memory=None,
                                steps=[('tfid_vect',
                                        TfidfVectorizer(analyzer='word',
                                                        binary=False,
                                                        decode_error='strict',
                                                        dtype=<class 'numpy.float64'>,
                                                        encoding='utf-8',
                                                        input='content',
                                                        lowercase=True,
                                                        max_df=1.0,
                                                        max
```

<!-- grade_id: cell-d07134c6fe5d056e, points: 1, grade: true, solution: false -->
```python
# Visible Testing
prediction = gs.predict(["I wish dogs knew how to speak English."])[0]
assert prediction in df.stars.values, 'You gs object should be able to accept raw text within a list. Did you include a vectorizer in your pipeline?'
```

## Part 4: Topic Modeling

Let's find out what those yelp reviews are saying! :D

1. Estimate a LDA topic model of the review text
    - Set num_topics to `5`
    - Name your LDA model `lda`
2. Create 1-2 visualizations of the results
    - You can use the most important 3 words of a topic in relevant visualizations. Refer to yesterday's notebook to extract. 
3. In markdown, write 1-2 paragraphs of analysis on the results of your topic model

When you instantiate your LDA model, it should look like this: 

```python
lda = LdaModel(corpus=corpus,
               id2word=id2word,
               random_state=723812,
               num_topics = num_topics,
               passes=1
              )

```

__*Note*__: You can pass the DataFrame column of text reviews to gensim. You do not have to use a generator.

## Note about  pyLDAvis

**pyLDAvis** is the Topic modeling package that we used in class to visualize the topics that LDA generates for us.

You are welcomed to use pyLDAvis if you'd like for your visualization. However, **you MUST comment out the code that imports the package and the cell that generates the visualization before you submit your notebook to CodeGrade.** 

Although you should leave the print out of the visualization for graders to see (i.e. comment out the cell after you run it to create the viz). 

```python
from gensim import corpora
# Due to limited computationalresources on CodeGrader, use the non-multicore version of LDA 
from gensim.models.ldamodel import LdaModel
import gensim
import re
```

### 1. Estimate a LDA topic model of the review tex

<!-- grade_id: cell-66331a185ff52f15, points: 0, grade: false, solution: true -->
```python
# Remember to read the LDA docs for more information on the various class attirbutes and methods available to you
# in the LDA model: https://radimrehurek.com/gensim/models/ldamodel.html

# don't change this value 
num_topics = 5

# use tokenize function you created earlier to create tokens 
df['lemmas'] = df['text'].apply(tokenize)
# create a id2word object (hint: use corpora.Dictionary)
id2word = corpora.Dictionary(df['lemmas'] )
# create a corpus object (hint: id2word.doc2bow)
corpus =  [id2word.doc2bow(doc_lemmas) for doc_lemmas in df['lemmas']]
# instantiate an lda model
lda = LdaModel(corpus=corpus,
               id2word=id2word,
               random_state=723812,
               num_topics = num_topics,
               passes=1
              )
# YOUR CODE HERE
```

#### Testing

<!-- grade_id: cell-5a3c181311134fa9, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert lda.get_topics().shape[0] == 5, 'Did your model complete its training? Did you set num_topics to 5?'
```

#### 2. Create 1-2 visualizations of the results

<!-- grade_id: cell-9b043e992fbd218c, points: 0, grade: false, solution: true -->
```python
# # import seaborn as sns
# # import matplotlib.pyplot as plt
# import pyLDAvis
# import pyLDAvis.gensim
# # Use pyLDAvis (or a ploting tool of your choice) to visualize your results 

# # YOUR CODE HERE
# pyLDAvis.enable_notebook()
# vis = pyLDAvis.gensim.prepare(lda, corpus, id2word)
# vis
```

#### 3. In markdown, write 1-2 paragraphs of analysis on the results of your topic model

```python
import spacy
from spacy.tokenizer import Tokenizer

nlp = spacy.load("en_core_web_lg")
tokenizer = Tokenizer(nlp.vocab)

sample_text = "Natural Language Processing is really fun!"
[token.text for token in tokenizer(sample_text)]
```

Output:
```
['Natural', 'Language', 'Processing', 'is', 'really', 'fun!']
```

```python

```